```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Lab 02: Numbers in Python

Welcome to Lab 02! Throughout the course you will complete a lab assignments like this one. You can't learn technical subjects without hands-on practice, so labs are an important part of the course.

Collaborating on labs is more than okay -- it's encouraged. You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** just copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

In today's lab, you'll learn about:

- Python integers and decimal numbers

- how Python stores integers and decimal numbers in memory.

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** Friday, February 11, 2022 at 11:59 pm

**Collaboration Policy:** Labs are a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

## 1. Integers

Python uses the class `int` to represent all integer numbers. All integers are objects. In general, classes and objects are basic building blocks in object-oriented programming languages. A class is written by a programmer to create an object. The class defines a set of properties and methods that are common to all objects of one type.

To store integers, a computer needs to represent the value using a base 2 number. For example, the number 5 is 

$$5=1 \times 2^2 + 0 \times 2^1 + 1 \times 2^0$$

As you can see it takes 3 bits to store the number 5 in memory:

$$(101)_2=(5)_{10}$$

Python doesn’t use a fixed number of bit to store integers. Instead, Python uses a variable number of bits to store integers. For example, 8 bits, 16 bits, 32 bits, 64 bits, 128 bits, and so on.

The maximum integer number that Python can represent depends on the memory available.

The following cell defines a variable that references an integer and uses the `type()` function to get the class name of the integer:

```python
three = 3
print(type(number))
```

### Getting the Size of an Integer

The `getsizeof()` function returns the number of bytes that Python uses to represent an integer. For example:

```python
from sys import getsizeof

size_of_2 = getsizeof(three)
print('To store the number', three, 'Python uses', size_of_2, 'bytes.')
```

**Question 1.** If Python uses 28 bytes to store a number, then how many bits are being used. Store this in the variable `bits_28`.

**Note:** 1 byte equals 8 bits.

```python
bits_28 = ...
bits_28
```

```python
grader.check("q1")
```

Let's take a look at the number 0.

```python
zero = 0

size_of_0 = getsizeof(zero)
print('To store the number', zero, 'Python uses', size_of_0, 'bytes.')
```

To store the number 0 Python only needs 1 bit, $\left(2^0 \right)$. Therefore, you can think that Python uses 24 bytes as an overhead for storing an integer object. The following cell returns the size of the integer 100:

```python
one_hundred = 100

size_of_100 = getsizeof(three)
print('To store the number', one_hundred, 'Python uses', size_of_100, 'bytes.')
```

**Question 2.** If we think of 24 bytes as overhead, how many bytes is Python really using to represent the number 100. Store this value in `overhead`.

```python
overhead = ...
overhead
```

```python
grader.check("q2")
```

Let's look at an larger integer, say $2^{10240}$. How many bytes do you think Python will use to store this value. Enter your choice in the next cell.

**Question 3.** Choose one of the following and enter your choice in the code cell below.

1. More than or equal 28 but less than or equal to 32

2. More than 32 but less than or equal to 36 

3. More than 36 but less than or equal to 48

```python
q3 = ...
```

```python
grader.check("q3")
```

Let's check.

```python
two_to_the_power_of_10240 = 2**10240

size_of_2_to_the_power_of_10240 = getsizeof(three)
print('To store the number', two_to_the_power_of_10240, 'Python uses', size_of_2_to_the_power_of_10240, 'bytes.')
```

Python integers support all standard operations including:

 * Addition `+`

* Subtraction `–`

* Multiplication `*`

* Division `/

* Exponents `**`

The result of addition, subtraction, multiplication, and exponents of integers is an integer. For example:

```python
a = 10
b = 20

c = a + b
print(c)
print(type(c))


c = a - b
print(c)
print(type(c))


c = a * b
print(c)
print(type(c))

c = a ** b
print(c)
print(type(c))
```

However, the division of two integers always returns a floating-point number. For example:

```python
a = 10
b = 5
c = a / b

print(c)
print(type(c))
```

## 2. Floats

Python uses the `float` class to represent the real numbers.

```python
point_1 = 0.1

print(point_1)
print(type(point_1))
```

For the sake of simplicity, significant digits are all digits except leading and trailing zeros. For example, 0.25 has two significant digits, 0.125 has three significant digits, and 12.25 has four significant digits.

$$(1.25)_{10} = \left(1 \times 2^0 + 0 \times 2^{-1} + 1 \times 2^{-2} \right)_{10} = (1.01)_2$$

Some numbers have a finite binary representation, but some don’t, e.g., $0.1$. It’s 

$$01.0001100110011 \ldots$$ in binary.

Because of this, Python can only use approximate float representations for those numbers.

Internally, Python can only represent 0.1 approximately.

To see how Python represents the 0.1 internally, you can use the `format()` function.

The following shows how Python represents the number 0.1 using 20 digits:

```python
format(0.1, '.20f')
```

Because Python can represent some floats approximately, it will cause many problems when you compare two floating-point numbers. As you can see, 0.1 is not exactly 0.1 but 

$$0.10000000000000000555...$$


**Question 4.** Will the following python code evaluate to True or False.

```
x = 0.1 + 0.1 + 0.1
y = 0.3

print(x == y)
```

Assign the boolean value of `True` or `False` to `q4`.

```python
q4 = ...
```

```python
grader.check("q4")
```

Internally, Python cannot use a finite number of digits to represent the numbers `x` and `y`:

```python
x = 0.1 + 0.1 + 0.1
y = 0.3

print(format(x, '.20f'))
print(format(y, '.20f'))
```

Note that the number of digits is infinite. We just show the first 20 digits.

One way to work around this problem is to round both sides of the equality expression to a number of significant digits. For example:

```python
print(round(x, 3))
print(round(y, 3))

print(round(x, 3) == round(y, 3))
```

This workaround doesn’t work in all cases.

[PEP485](https://www.python.org/dev/peps/pep-0485/) provides a solution that fixes this problem by using relative and absolute tolerances.

It provides the `isclose()` function from the math module returns True if two numbers are relatively close to each other.

The following shows the `isclose()` function signature:

```
isclose(a, b, rel_tol=1e-9, abs_tol=0.0)
```

```python
from math import isclose

print(isclose(x,y))
```

<!-- BEGIN QUESTION -->

**Question 5.** Write a short summary of the difference between `int` and `float` in Python. Include a description of possible things to consider when performing computations using each class.

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: None)
```
>>> bits_28
224
```


Tests for q2 (points: None)
```
>>> overhead
4
```


Tests for q3 (points: None)
```
>>> isinstance(q3, int)
True
```
```
>>> q3 in [1, 2, 3]
True
```


Tests for q4 (points: None)
```
>>> isinstance(q4, bool)
True
```
```
>>> q4
False
```
