```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

# Lab 01: Collection Data Types and Arrays

Welcome to Lab 01! Throughout the course you will complete a lab assignments like this one. You can't learn technical subjects without hands-on practice, so labs are an important part of the course.

Collaborating on labs is more than okay -- it's encouraged. You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** just copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

In today's lab, you'll learn how to:

- import Python libraries.

- write user-defined functions.

- plot graphs.

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** Friday, February 4, 2022 at 11:59 pm

**Collaboration Policy:** Labs are a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

## 1. Keyboard Shortcuts

Even if you are familiar with Jupyter, we strongly encourage you to become proficient with keyboard shortcuts (this will save you time in the future). To learn about keyboard shortcuts, go to **Help $\to$ Keyboard Shortcuts** in the menu above.

Here are a few that we like:

1. Ctrl + Return: Evaluate the current cell
2. Shift + Return: Evaluate the current cell and move to the next
3. ESC: command mode (may need to press before using any of the commands below)
 -  a: create a cell above
 -  b: create a cell below
 - dd: delete a cell
 -  z: undo the last cell operation
 -  m: convert a cell to markdown
 -  y: convert a cell to code

## 2. Importing Libraries and Magic Commands

In Numerical Analysis, we will be using common Python libraries to help us perform mathematical operations. A library is an umbrella term referring to a reusable chunk of code. Usually, a Python library contains a collection of related modules and packages. Actually, this term is often used interchangeably with “Python package” because packages can also contain modules and other packages (subpackages). However, it is often assumed that while a package is a collection of modules, a library is a collection of packages.

By convention, we import all libraries at the very top of the notebook. There are also a set of standard aliases that are used to shorten the library names. Below are some of the libraries that you may encounter throughout the course, along with their respective aliases.

### Importing Libraries

We'll look at the `math` module as a first example. The `math` module is extremely useful in computing mathematical expressions in Python. 

Suppose we want to very accurately compute the area of a circle with a radius of 5 meters.  For that, we need the constant $\pi$, which is roughly 3.14.  Conveniently, the `math` module has `pi` defined for us.

Run the cell below, but please **do not** change it.

```python
import math

radius = 5
area_of_circle = radius**2*math.pi
area_of_circle
```

In the code above, the line `import math` imports the math module. This statement creates a module and then assigns the name `math` to that module. We are now able to access any variables or functions defined within `math` by typing the name of the module followed by a dot, then followed by the name of the variable or function we want.

**Question 1.** The module `math` also provides the name `e` for the base of the natural logarithm, which is roughly 2.71.  Compute $e^{\pi}-\pi$, giving it the name `near_twenty`.

```python
near_twenty = ...
near_twenty
```

```python
grader.check("q1")
```

In the question above, you accessed variables within the `math` module. 

Modules also define functions.  For example, `math` provides the name `sin` for the sine function.  Having imported `math` already, we can write `math.sin(3)` to compute the sine of 3.

**Note:** This sine function considers its argument to be in [radians](https://en.wikipedia.org/wiki/Radian), not degrees. 180 degrees are equivalent to $\pi$ radians.

**Question 2.** A $\frac{\pi}{4}$ radian (45-degree) angle forms a right triangle with equal base and height, pictured below.  If the hypotenuse (the radius of the circle in the picture) is 1, then the height is $\sin\left(\frac{\pi}{4}\right)$.  Compute that value using `sin` and `pi` from the `math` module.  Give the result the name `sine_of_pi_over_four`.

<img src="http://mathworld.wolfram.com/images/eps-gif/TrigonometryAnglesPi4_1000.gif">

**Source:** [Wolfram MathWorld](http://mathworld.wolfram.com/images/eps-gif/TrigonometryAnglesPi4_1000.gif)

```python
sine_of_pi_over_four = ...
sine_of_pi_over_four
```

```python
grader.check("q2")
```

For your reference, the cells below demonstrate some more examples of functions from the `math` module.

Notice how different functions take in different numbers of arguments. Often, the [documentation](https://docs.python.org/3/library/math.html) of the module will provide information on how many arguments are required for each function.

Run the cell below.

```python
# Calculating logarithms (the logarithm of 8 in base 2)
# The result is 3 because 2 to the power of 3 is 8
math.log(8, 2)
```

Run the cell below.

```python
# Calculating square roots
math.sqrt(5)
```

There are various ways to import and access code from outside sources. The method we used above — `import <module_name>` — imports the entire module and requires that we use `<module_name>.<name>` to access its code. 

We can also import a specific constant or function instead of the entire module. Notice that you don't have to use the module name beforehand to reference that particular value. However, you do have to be careful about reassigning the names of the constants or functions to other values.

Run the cell below.

```python
# Importing just cos and pi from math
# We don't have to use `math.` in front of cos or pi

from math import cos, pi
print(cos(pi))

# We do have to use it in front of other functions from math, though
math.log(pi)
```

Or we can import every function and value from the entire module.

Run the cell below.

```python
# Lastly, we can import everything from math using the *
# Once again, we don't have to use 'math.' beforehand 

from math import *
log(pi)
```

Don't worry too much about which type of import to use. It's often a coding style choice left up to each programmer. In this course, you'll always import the necessary modules when you run the setup cell (like the first code cell in this lab).

### Magic Commands

`%matplotlib inline` is a Jupyter magic command that configures the notebook so that Matplotlib displays any plots that you draw directly in the notebook rather than to a file, allowing you to view the plots upon executing your code.

Run the cell below.

```python
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline
```

Another useful magic command is `%%time`, which times the execution of that cell. You can use this by writing it as the first line of a cell.

**Note:** `%%` is used for cell magic commands that apply to the entire cell, whereas `%` is used for line magic commands that only apply to a single line.

```python
%%time

# Initialize an empty list
my_list = []

for i in range(100):  # for loop to iterate through the values 0 to 99
    my_list.append(i) # The .append method will add an item to the end of a list 
```

## 2. Collections of Elements

The cell above uses a list, a `for` loop, the `range` function, and the `.append()` method.

#### Lists 

* Lists are used to store multiple items in a single variable. 

* Lists are a built-in data type in Python used to store collections of data.
.
* Lists may contain different types

* List items are ordered, changeable, and allow duplicate values. 

* List items are indexed, the first item has index `[0]`, the second item has index `[1]` etc

* Lists are created using square brackets. For example, `my_list = [2, 3, 5, 7, 11]`, is a list of the first 5 prime numbers.

#### `for` loop

* A `for` loop is used for iterating over a sequence (like a list).

* A `for` loop will execute a set of statements, once for each item in a sequence (like a list).

#### `range`

* The `range()` function returns a sequence of numbers, starting from 0 by default, and increments by 1 (by default), and stops before a specified number. The syntax for the `range()` function is `range(start, stop, step)`.

#### `.append()`

* The `.append()` method will "append" an item to the end of a list (we will cover this method in more detail later).

**Question 3.** Use the `range()` function to generate a list of the first 100 odd integers. Save the list of numbers to `first100_odd_integers`.

**Hint:** Review the syntax of the `range()` function. 

```python
first100_odd_integers = ...
```

```python
grader.check("q3")
```

The output from **Question 3.** looked like this

```
range(1, 201, 2)
```

When we ran the code cell the variable name `first100_odd_integers` was on the last line. We didn't see the list of numbers, only what was on the right of the equal sign. If we want to see actual list of numbers we can use a `for` loop.

Run the cell below.

```python
for number in range(1, 201, 2):
    print(number)
```

`for` loops are used when you have a block of code which you want to repeat a fixed number of times. The `for` loop is always used in combination with an iterable object, like a *list* or a *range*. The Python `for` statement iterates over the members of a sequence in order, executing the block each time.

In the `for` loop in the previous code cell the name `number` is used as a container for each element in the range of odd integers from 1 to 199. The value of `number` is updated each time the loop iterates over the range of odd integers. The choice of the name `number` is nothing special. We could have used any acceptable Python name. For example, we could have written

```
for i in range(1, 201, 2):
    print(i)
```

As long as the name is valid meaning it only consists of:

* uppercase and lowercase letters ( A-Z , a-z )

* digits ( 0-9 )

* the underscore character ( _ )

and 

* doesn't start with a digit (0-9) 

we can use it.

## 3. `Numpy`

It seems that is would be beneficial to be able to view our list of values without having to resort to writing a `for` loop.

In our class we will do a lot computations and plotting. Thus, we need to be able to work with lists of numbers. To work with lists of numbers we will use the `NumPy` library. 

#### What is `NumPy`?

`NumPy` is a library for the Python programming language, adding support for large, multi-dimensional arrays and matrices, along with a large collection of high-level mathematical functions to operate on these arrays. `NumPy` (Numerical Python) is an open source Python library that’s used in almost every field of science and engineering.  Click [here](https://numpy.org/doc/stable/user/absolute_beginners.html) to read more about the  basics for beginners. 

Let's import the library by running the cell below.

```python
import numpy as np
```

Let's make an array of numbers from 1 to 10 including the endpoints.

```python
np.arange(1, 11)
```

Let's make an array of numbers from 0.1 to 1.0 including the endpoints.

```python
np.arange(0.1, 1.1, 0.1)
```

To use the `arange` function from the `NumPy` library need to put the `np` alias in front of the function name. The syntax for the `np.arange` function is similar to the `range` function.

```
arange([start,] stop[, step,][, dtype])
```

* start : [optional] start of interval range. By default start=0

* stop  : end of interval range

* step  : [optional] step size of interval. By default step size=1.

* dtype : type of output array

**Question 4.** Use the `np.arange()` function to generate a list of the first 100 even integers. Save the list of numbers to `first100_even_integers`.

**Hint:** Review the syntax of the `np.arange()` function. 

```python
first100_even_integers = ...
first100_even_integers
```

```python
grader.check("q4")
```

Notice how the values get displayed when we run the code cell. 

Another great benefit of an array is the ability to do *element-wise* operations. Run the next three cells to see examples.

```python
# Create an array of numbers from 0 to 200
numbers = np.arange(1, 201)

# Divide each element by 100
numbers/100
```

```python
# Take the log base 10 of each element
np.log10(numbers)
```

```python
# Take the squre root of each element
numbers**(1/2)
```

**Question 5.**  Use `np.arange` and the exponentiation operator `**` to compute the first 30 powers of 2, **starting from $2^0$**. Save this to `powers_of_2`.

**Hint:** Part of your solution will involve `np.arange`, but your array shouldn't have more than 30 elements.

**Warning:** `np.arange(1, 2**30, 1)` creates an array with $2^{30}$ elements and **will crash your kernel** and you will need to start over. 

```python
powers_of_2 = ...
powers_of_2
```

```python
grader.check("q5")
```

## 4. Plotting Graphs

### Matplotlib

Take time to go through the `pyplot` tutorial. Please click [here](https://matplotlib.org/stable/tutorials/introductory/pyplot.htmlhttps://matplotlib.org/stable/tutorials/introductory/pyplot.html) and go through the tutorial notebook and familiarize yourself with the basics of `pyplot`. This should take roughly 20 minutes. We've already imported the `matplotlib` library earlier in the notebook.

```
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline
```

**Note:** The tutorial uses `np.arange`, which returns an array that steps from $a$ to $b$ with a fixed step size $s$. While this is fine in some cases, we sometimes prefer to use `np.linspace(a, b, N)`, which divides the interval $[a, b]$ into $N$ equally spaced points.

For example, `np.linspace` always includes both end points while `np.arange` will not include the second end point $b$. For this reason, when we are plotting ranges of values we tend to prefer `np.linspace`.

Notice how the following two statements have different parameters but return the same result.

```python
np.arange(-5, 6, 1.0)
```

```python
np.linspace(-5, 5, 11)
```

<!-- BEGIN QUESTION -->

**Question 6.** Let's visualize the function $f(t) = 3\sin(2\pi t)$.

- Set the $x$ limit of all figures to $[0, \pi]$ and the $y$ limit to $[-10, 10]$. 

- Plot the sine function using `plt.plot` with 30 red plus signs. 

- Make sure the $x$ ticks are labeled $\left[0, \frac{\pi}{2}, \pi\right]$, and that your axes are labeled as well. 

**Note:** Since this question will be manually graded, there is no `grader.check("q7")`.


**Documentation:** Click [here](https://matplotlib.org/api/pyplot_api.html) to use the `matplotlib` documentation for reference.


Your plot should look like the following:

<center><img src="images/graph1.png"></center>

**Hints:** 

* You can set axis bounds with `plt.axis`.

* You can set `xticks` and labels with `plt.xticks`.

* Make sure you add `plt.xlabel`, `plt.ylabel`, and `plt.title`.

```python
t = ...
y = 3*np.sin(2*math.pi*t)

plt.ylim((..., ...))
plt.scatter(t, y, marker = ..., c = ..., linewidth=1)
plt.xticks([0, math.pi/2, math.pi],[r'$0$', r'$\pi/2$', r'$\pi$'])
plt.xlabel('t')
plt.ylabel('f(t)')
plt.title('f(t) = 3sin(2$\pi$t)');
```

<!-- END QUESTION -->

## 5. Small Numbers

### How Python Treats Numbers

Python has three built-in numeric data types: integers, floating-point numbers, and complex numbers. Today we will only be concerned with integers and floating-point numbers. Integers and floating points are separated by the presence or absence of a decimal point. For instance, 5 is an integer whereas 5.0 is a floating-point number. Run the cells below to see examples.

```python
5
```

```python
5.0
```

We can use the `type` function to know which type of a value or a variable belongs. 

```python
print(type(5))
print(type(5.0))
```

```python
x1 = 5
x2 = 5.0

print(type(x1))
print(type(x2))
```

A central them of this course is the concept of approximations. Python's built-in floating-point class performs some calculations that might we might not agree with. We all know that the sum of `1.1` and `2.2` is `3.3`, but Python seems to disagree. 

Run the cell below to see.

```python
(1.1+2.2)  == 3.3
```

Why does Python think this is **False**? Run the cell below to see what Python thinks. 

```python
1.2+2.2
```

To quote the [documentation](https://docs.python.org/3.3/tutorial/floatingpoint.html):

> Unfortunately, most decimal fractions cannot be represented exactly as binary fractions. A consequence is that, in general, the decimal floating-point numbers you enter are only approximated by the binary floating-point numbers actually stored in the machine.

Sometimes we will get what we expect. Run the cells below to see some examples.

```python
1.1+1.1
```

```python
1.1+2.3
```

```python
1.1+2.2
```

Finally, below is another weird one.

```python
0.1+0.1+0.1-0.3
```

```python
99.999999999-99.999999997
```

The `e-09` means $10^{-9}$. That's kind of a small number. 

What we want to do is to find the smallest number (not 0) Python treats equal to 0. For example, 

$$1 + \text{a really small number that's not 0} = 1$$

**Question 7.** Assign a value to `a_really_small_number_thats_not_0`, then run the cell. Once you get 1.0 you are done with the question.

```python
a_really_small_number_thats_not_0 = ...
1 + a_really_small_number_thats_not_0
```

```python
grader.check("q7")
```

The next natural question has to be, what number (not 0) does Python say is equal to zero.

<!-- BEGIN QUESTION -->

**Question 8.** Assign a value to `not_really_0`, then run the cell. Once the result of running the code cell is **True** you are done. Be careful, we are looking for the **first** number that Python says is equal to 0. Hence typing `1e-500` won't be an acceptable answer.

**Note:** Since this question will be manually graded, there is no `grader.check("q9")`.

```python
not_really_0 = ...
not_really_0 == 0.0
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: None)
```
>>> near_twenty
19.99909997918947
```


Tests for q2 (points: None)
```
>>> sine_of_pi_over_four
0.7071067811865475
```


Tests for q3 (points: None)
```
>>> isinstance(first100_odd_integers, range)
True
```
```
>>> [199, 1, 100] == [max(first100_odd_integers), min(first100_odd_integers), len(first100_odd_integers)]
True
```


Tests for q4 (points: None)
```
>>> isinstance(first100_even_integers, np.ndarray)
True
```
```
>>> [200, 2, 100] == [max(first100_even_integers), min(first100_even_integers), len(first100_even_integers)]
True
```


Tests for q5 (points: None)
```
>>> isinstance(powers_of_2, np.ndarray)
True
```
```
>>> np.sum(powers_of_2)
1073741823
```


Tests for q7 (points: None)
```
>>> np.random.seed(123)
>>> np.random.randint(1e10)*a_really_small_number_thats_not_0
2.967327842e-07
```
