```python
# Initialize Otter
import otter
grader = otter.Notebook("lab11.ipynb")
```

# Lab 11: More Numerical Integration

Welcome to Lab 11! Throughout the course you will complete a lab assignments like this one. You can't learn technical subjects without hands-on practice, so labs are an important part of the course. 

Collaborating on labs is more than okay -- it's encouraged. You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too $-$ the best way to solidify your knowledge of a subject is to explain it. You should not just copy someone else's answers, but rather work together to gain understanding of the task you need to complete.

In today's lab, you will be asked to use your python programming skills to write code to approximate definite integrals.

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** Tuesday, May 10, 2022 at 11:59 pm

**Collaboration Policy:** Labs are a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

Run the cell below to import the required modules.

```python
import numpy as np
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')
%matplotlib inline
```

Before we get started let's review two things:

- `np.linspace()`

- `lambda` functions

# 0. Review

The `np.linspace()` function returns evenly spaced numbers over a specified interval. For example, 

`np.linspace(0, 10, 11)` 

would return 

`array([ 0.,  1.,  2.,  3.,  4.,  5.,  6.,  7.,  8.,  9., 10.])`. 

Since this is an array object we can iterate over each element. Remember, Python gives the index value 0 for the first element. For example, if we saved our array to `x`

`x = np.linspace(0, 10, 11)` 

then ran the command 

`x[1]`

we would get 1. If instead we wanted the first element, then we would need to run `x[0]`.

Try it for yourself.

```python
np.linspace(0, 10, 11)
```

```python
x = np.linspace(0, 10, 11)
x[1]
```

```python
x[0]
```

If we wanted to access elements 1, 3, 5, 7, and 9 we could slice the array

```
x[1::2]
```

or, if we wanted to access elements 2, 4, 6, and 8 

```
x[2::2]
```

Try it for yourself.

```python
x[1::2]
```

```python
x[2::2]
```

Let's examine what happens when we run the command

```
x[1::2]
```

Slicing in Python means taking elements from one given index to another given index.

- We pass slice instead of index like this: `[start:end]`.

- We can also define the step, like this: `[start:end:step]`.

- If we don't pass start its considered 0.

- If we don't pass end its considered length of array in that dimension.

- If we don't pass step its considered 1.

So our command 

```
x[1::2]
```

means to start at the **second** element in the array, end at the next to the last element, and use a step size of 2.

**Note:** The end will not be included in the slice. For example, `x[:2]` would return `array([0., 1.])`. If we wanted to include 2, then we would need to run `x[:3]`.



We can use an array to provide the inputs for a `lambda` function. Remember the syntax for a `lambda` function is 

```
f = lambda x: x**2
```

If we define `f` we input the values from our array into `f`


```
f([x[1::2])
```

Try it and see.

```python
x[:2]
```

```python
x[:3]
```

```python
f = lambda x: x**2

f(x[1::2])
```

Throughout this notebook you can use `np.linspace` and `lambda` functions to help you complete the exercises. 

**Note:** There are other methods that can be used to accomplish the same task. Feel free to use other methods, if you have experience using Python. 

# 1. Graphing 

Let's write a lambda functions that we can use to graph each function over its interval of integration.

<!-- BEGIN QUESTION -->

**Question 1.** Write a lambda function to sketch the graph of $f(x)=\sin x$.

**Hint:** Use sine from numpy (e.g., `np.sin`)

```python
x  = np.linspace(0, np.pi/2, 100)
f = ...
plt.plot(x, f(x), color = 'black', ms = 3)
plt.title("sin(x)");
```

<!-- END QUESTION -->

# 2. Simpson's Rule

Simpson's Rule (named after the English mathematician Thomas Simpson 1710−1761) is a numerical method that approximates the value of a definite integral by using quadratic functions. The rule is based on the fact that given three points, we can find the equation of a quadratic through those points. To obtain an approximation of the definite integral 

$$\int_a^b f(x) \ dx$$
 
using Simpson's Rule, we partition the interval $[a,b]$ into an even number $n$ of subintervals, each of width

$$\Delta x = \frac{b-a}{n}$$

On each pair of consecutive subintervals $\left[x_{i-1}, x_i \right], \left[x_i, x_{i+1} \right]$ we consider a quadratic function $f(x)=ax^2+bx+c$ such that it passes through the points $\left(x_{i-1}, f\left(x_{i-1} \right) \right), \left(x_i, f\left(x_i \right) \right), \left(x_{i+1}, f\left(x_{i+1} \right) \right)$.


<center><img src="images/simpsons-rule1.svg" width="300" height="400"></center>

If the function $f(x)$ is continuous on $[a,b]$, then 

$$\int_a^b f(x) \  dx \approx \frac{\Delta x}{3}\left[f\left(x_0 \right)+4f\left(x_1 \right)+2f\left(x_2 \right)+4f\left(x_3 \right) + 2f\left(x_4 \right) + \cdots +4f\left(x_{n-1} \right) + f\left(x_n \right) \right]$$

The coefficients in Simpon's Rule have the following pattern $1,4,2,4,2,\ldots,4,2,4,1$.

Our goal is to write a Python function to approximate the value of a definite integral using Simpson's Rule.

<!-- BEGIN QUESTION -->

**Question 2.** Write a function which takes input parameters f, a, b and n and returns the approximation of the definite integral of of $f(x)$ on the interval $[a, b]$. Furthermore, let's assign a default value of $n=50$.

```python
def simpsons_rule(f, a, b, n = 50):
    """
    Purpose
    -------
    Approximate the integral of f(x) from a to b by Simpson's rule.

    Parameters
    ----------
    f : lambda function of a single variable
    a , b : Interval of integration [a,b]
    n : Number of subintervals of [a,b] (n is an even integer)

    Returns
    -------
    Approximation of the integral of f(x) from a to b using Simpson's rule with n subintervals of equal length.

    Example
    --------
    >>> simpsons_rule(lambda x : np.sin(x), 0, np.pi, 10)
    2.0001095173150043
    """
    if n % 2 == 1:
        raise ValueError("n must be an even integer.")
    ...
```

<!-- END QUESTION -->

Run the cell below to make sure your function is correct. It should return `2.0001095173150043`.

```python
simpsons_rule(lambda x : np.sin(x), 0, np.pi, 10)
```

<!-- BEGIN QUESTION -->

**Question 3.** Use your function to approximate the value of $\displaystyle \int_1^{1.6} \frac{2x}{x^2-4} \ dx$.

```python
f = ...
simpsons_rule(f, 1, 1.6)
```

<!-- END QUESTION -->

Suppose a radar gun was used to record the speed of a runner during the first 5 seconds of a race (see the table). 

|t(s)|v(m/s)|
|----|------|
|0|0|
|0.5|4.67|
|1.0|7.34|
|1.5|8.86|
|2.0|9.73|
|2.5|10.22|
|3.0|10.51|
|3.5|10.67|
|4.0|10.76|
|4.5|10.81|
|5.0|10.81|

**Question 4.** Make a `numpy` array of the values fro v(m/s) column.

**Hint:** Click [here](https://numpy.org/doc/stable/user/basics.creation.html) to see the syntax on creating a `numpy` array.

```python
v = ...
v
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

**Question 5.** Use Simpson’s Rule to estimate the distance the runner covered during those 5 seconds.

```python
...
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q4 (points: None)
```
>>> len(v)
11
```
```
>>> np.sum(v)
94.38000000000001
```
