```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

# Lab 04: Root Finding

Welcome to Lab 04! Throughout the course you will complete a lab assignments like this one. You can't learn technical subjects without hands-on practice, so labs are an important part of the course.

Collaborating on labs is more than okay -- it's encouraged. You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** just copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

In today's lab, you'll learn about root finding algorithms. 

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** Thursday, February 24, 2022 at 11:59 pm

**Collaboration Policy:** Labs are a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

## 0. Python Functions

Let's write a very simple function that converts a proportion to a percentage by multiplying it by 100.  For example, the value of `to_percentage(.5)` should be the number 50 (no percent sign).

A function definition has a few parts.

##### `def`
It always starts with `def` (short for **def**ine):

    def

##### Name
Next comes the name of the function.  Like other names we've defined, it can't start with a number or contain spaces. Let's call our function `to_percentage`:
    
    def to_percentage

##### Signature
Next comes something called the *signature* of the function.  This tells Python how many arguments your function should have, and what names you'll use to refer to those arguments in the function's code.  A function can have any number of arguments (including 0!). 

`to_percentage` should take one argument, and we'll call that argument `proportion` since it should be a proportion.

    def to_percentage(proportion)
    
If we want our function to take more than one argument, we add a comma between each argument name. Note that if we had zero arguments, we'd still place the parentheses () after than name. 

We put a colon after the signature to tell Python it's over. If you're getting a syntax error after defining a function, check to make sure you remembered the colon!

    def to_percentage(proportion):

##### Documentation
Functions can do complicated things, so you should write an explanation of what your function does.  For small functions, this is less important, but it's a good habit to learn from the start.  Conventionally, Python functions are documented by writing an **indented** triple-quoted string:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""
    
    
##### Body
Now we start writing code that runs when the function is called.  This is called the *body* of the function and every line **must be indented with a tab**.  Any lines that are *not* indented and left-aligned with the def statement is considered outside the function. 

Some notes about the body of the function:
- We can write code that we would write anywhere else.  
- We use the arguments defined in the function signature. We can do this because we assume that when we call the function, values are already assigned to those arguments.
- We generally avoid referencing variables defined *outside* the function. If you would like to reference variables outside of the function, pass them through as arguments!


Now, let's give a name to the number we multiply a proportion by to get a percentage:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""
        factor = 100

##### `return`
The special instruction `return` is part of the function's body and tells Python to make the value of the function call equal to whatever comes right after `return`.  We want the value of `to_percentage(.5)` to be the proportion .5 times the factor 100, so we write:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""
        factor = 100
        return proportion * factor
        
`return` only makes sense in the context of a function, and **can never be used outside of a function**. `return` is always the last line of the function because Python stops executing the body of a function once it hits a `return` statement.

**Note:**  `return` inside a function tells Python what value the function evaluates to. However, there are other functions, like `print`, that have no `return` value. For example, `print` simply prints a certain value out to the console. Be aware that `return` and `print` are **very** different. 

**Question 1.** Define `to_percentage` in the cell below.  Call your function to convert the proportion $0.2$ to a percentage.  Name that percentage `twenty_percent`.


```python
def ...
    ``` ... ```
    ... = ...
    return ...
twenty_percent = ...
twenty_percent
```

```python
grader.check("q1")
```

Like you've done with built-in functions in previous labs (`max`, `abs`, etc.), you can pass in named values as arguments to your function.

**Question 2.** Use `to_percentage` again to convert the proportion named `a_proportion` (defined below) to a percentage called `a_percentage`.

**Note:** You don't need to define `to_percentage` again. Like other named values, functions stick around after you define them.

```python
a_proportion = 2**(.5)/2
a_percentage = ...
a_percentage
```

```python
grader.check("q2")
```

Here's something important about functions: the names assigned *within* a function body are only accessible within the function body. Once the function has returned, those names are gone.  So even if you created a variable called `factor` and defined `factor = 100` inside of the body of the `to_percentage` function and then called `to_percentage`, `factor` would not have a value assigned to it outside of the body of `to_percentage`. 

Run the cell below. You should see an error (if you don't, you might have defined `factor` somewhere above).

```python
factor
```

## 1. Introduction to Numerical Root Finding

The topic of numerical root finding really boils down to approximating the solutions to equations without using all of the by-hand techniques that you've learned up to this point in high school. The down side to everything that we’re about to do is that our answers are only ever going to be approximations. Mathematicians love (and I mean love) exact values. Think about values like $\pi$, $e$, $\sqrt 2$ and all those values for sine and cosine you may have been forced to memorize from special right triangles and the unit circle. 

There are two main reasons we would want to do numerical algebra: 

1. Algebra by hand can be extremely challenging, time consuming, and error prone.

2. The vast majority of equations do not lend themselves to by-hand solutions. 


### The Quadratic Formula

We will start with quadratic formula 

$$x=\frac{-b \pm \sqrt{b^2-4ac}}{2a}$$

The quadratic formula can be used to find the roots (i.e. zeros) of a quadratic equation of the form $ax^2+bx+c=0$ where $a$, $b$, and $c$ are real numbers and $a \ne 0$. If we want to write a function that returns the real roots of a quadratic equation we need to consider the following cases:

- If $b^2 < 4ac$, then roots are complex (not real). For example, roots of $x^2 + x + 1$ are $\frac{-1 + i\sqrt{3}}{2}$ and $\frac{-0.5 - i\sqrt{3}}{2}$.

- If $b^2=4ac$, then the roots are real and equal to each other. For example, root of $x^2-4x+4$ is $2$.

- If $b^2 > 4ac$, then roots are real and different. For example, roots of $x^2-4x-5$ are $5$ and $-1$.

**Note:** To write our function to find the roots we will import the `math` module.

```python
import math
```

**Question 3.** Write a function to find the real roots of a quadratic equation. 

**Note:** You can assume that the values for $a$, $b$, and $c$ are real numbers. But you need consider the following cases:

- If $b^2 < 4ac$, then print the statement "No real roots.".

- If $b^2=4ac$, then print the statement "One real root." and return the value of the root.

- If $b^2 > 4ac$, then print the statement "Two distinct real roots" and return both roots.

```python
def quadratic_roots(a,b,c): 
    """
    Parameters
    ----------
    a, b, and c: The coefficients of the quadratic equation
    
    Returns
    -------
    The values of the real roots of a quadratic equation
    """
...
quadratic_roots(1,10,25)
```

```python
grader.check("q3")
```

The quadratic formula is sufficient for quadratic equations, but what if our polynomial is not degree two. In class this week we learned two methods for approximating the roots of polynomials: the Bisection method and Newton's method.

## 2. The Bisection Method

#### Bisection Method

Previously in this lab we wrote a function to find the real roots of a quadratic equation. Now let's do the same thing for the Bisection method. From your notes in class you should have an idea of how the Bisection method works. Review your notes if you need to refresh your memory.

For our function we will use $a$ and $b$ for inputs for an interval in which $f(a)$ and $f(b)$ have opposite signs.

**Note:** You will need to define a function $f(x)$ that returns the values of $f(a)$ and $f(b)$. For example,

```
def practice_function(x):
    return (the expression for your function should go here)
```

You can use the empty code cell below to practice defining and graphing function. First import `NumPy` and `matplotlib`.

```python
import numpy as np
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')
%matplotlib inline
```

```python
def practice_function(x):
    return x**2-1

x = np.arange(-5,5.5,0.5)
plt.plot(x,practice_function(x), '-ok', ms = 3);
```

Let's say we want to graph the function Mr. Rash used during class to explain the bisection method:

$$y=x^6-x-1$$

<!-- BEGIN QUESTION -->

**Question 4.** Write a python function for $y=x^6-x-1$ and sketch the graph over the interval $[-1,1]$.

```python
def y(x):
   return ...

x = ...
...
```

<!-- END QUESTION -->

This function is continuous over the interval $[-1, -0.5]$ and the Intermediate Value Theorem can be applied on the interval $[-1, -0.5]$, so we should be able to use the Bisection method to find the root on the interval. 

Before we proceed, let's make sure we know what the Bisection method is actually doing.

<!-- BEGIN QUESTION -->

**Question 5.** Explain the how the Bisection method works.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Now let's implement the method and look at a table of values.

<!-- BEGIN QUESTION -->

**Question 6.** Run the function to implement the Bisection method on the function $f(x)=x^6-x-1$.

**Notes:** 

- If $f(a)$ and $f(b)$ do not have opposite signs stop the loop, print the message "A root is not guaranteed in this interval.".

- You **must** use a for loop and set the number of iterations to 25 (`iterations = 25`).

- The inputs should be the endpoints of the interval $a$ and $b$ as well as a level of tolerance (`tol = 1e-6`).

- If you do not reach the specified level of tolerance after 25 iterations then print the message "Level of tolerance not reached after 25 iterations." and return the value.

- If you reach the level of tolerance before 25 iterations stop the loop and return the value. A `break` statement can be used to end a loop. Click [here](https://www.geeksforgeeks.org/python-break-statement/) to read about how it is used.

```python
from scipy import optimize

def f(x):
    return x**6-x-1

def bisection_method(a, b, iterations = 25, tol = 1e-6):
    """
    Parameters
    ----------
    a, b:            The endpoints of the interval
    coefs:           The coefficients of the polynomial as a list to find the true value of the root on the interval
    iterations = 25: Set the maximum number of iterations for the loop
    tol = 1e-6:      Set the level of tolerance for the distance between x_i and the midpoint
    
    Returns
    -------
    The approximation for the root
    """
    x_is = []
    alpha = optimize.root_scalar(f, bracket=[-1,-0.5], method = 'brentq').root
    print("{:<8} {:<20} {:<30} {:<30}".format('i','alpha','x_i', '|alpha-x_i|'))
    
    if f(a)*f(b) >= 0:
        print("A root is not guaranteed in this interval.")
    
    for i in range(iterations):
        c = (a + b)/2.0
        x_is.append(c)
        print("{:<8} {:<20} {:<30} {:<30}".format(i, alpha, c, abs(alpha-x_is[i])))
                  
        if f(a)*f(c) > 0:
            a = c
        else:
            b = c
            
        if abs(b - a) < tol:
            break
        elif i == iterations-1:
            print("\nLevel of tolerance not reached after 25 iterations.")
            break
    
    return c 

print("\nAnswer:\t", bisection_method(-1, -0.5))
```

<!-- END QUESTION -->

#### Examining the Bisection Method

<!-- BEGIN QUESTION -->

**Question 7.** How many iterations do you think it took to reach the desired level of tolerance? Explain.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8.** Mention the strengths and weaknesses of this approximation technique.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Newton's Method

Mr. Rash also went over Newton's method. We learned that this approximation technique also has strengths and weaknesses.

<!-- BEGIN QUESTION -->

**Question 9.** Explain the how Newton's method works.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10.** Under what conditions would Newton's method fail to converge to a root.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Optional Programming Challenge

<!-- BEGIN QUESTION -->

**Question 11.** Write a python function to implement Newton's method.

```python
def newtons_method(...):
    ...
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: 0)
```
>>> twenty_percent
20.0
```


Tests for q2 (points: 0)
```
>>> a_percentage
70.71067811865476
```


Tests for q3 (points: None)
```
>>> quadratic_roots(1,1,4)
No real roots.
```
```
>>> quadratic_roots(1,-4,4)
One real root.
2.0
```
```
>>> quadratic_roots(1,10,4)
 Two distinct real roots.
(-0.41742430504416017, -9.582575694955839)
```
