```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

# Lab 05: Root Finding

Welcome to Lab 05! Throughout the course you will complete a lab assignments like this one. You can't learn technical subjects without hands-on practice, so labs are an important part of the course.

Collaborating on labs is more than okay -- it's encouraged. You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** just copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

In today's lab, you'll learn more about root finding algorithms. 

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** Tuesday, March 8, 2022 at 11:59 pm

**Collaboration Policy:** Labs are a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

## 0. Lambda Functions

Sometimes we want to write a small function or we want a function that we don't need to keep because we aren't going to use it later. In python we can write functions using a shorthand technique. These small, unnamed (anonymous) functions are called lambda functions. For example, suppose we want to square a number using a lambda function (I know it seems silly but humor me), we could use a lambda function.

```
square_biz = lambda x : x**2
```

The `square_biz` lambda function is written in one line of code. The argument is `x` and the `:` separates the argument form the expression. Notice that unlike the functions from the previous lab we don't need the keyword `def`. If we wanted to use our lambda function we would enter

```
square_biz(5)
```

and our function will return the value 25 (lambda functions can't use a `return` statement or some other keywords).

Let's try it.

**Question 1.** Write a **lambda** function named `f` that will compute the value of $f(x)=x^6-x-1$.

```python
f = ...
```

```python
grader.check("q1")
```

Perhaps we want to find values of the derivative for $f(x)=x^6-x-1$.

**Question 2.** Write a **lambda** function named `Df` that will compute the value of $f'(x)$, where $f(x)=x^6-x-1$..

```python
Df = ...
```

```python
grader.check("q2")
```

We can also plot lambda functions. run the cell below to load the `NumPy` and `matplotlib` modules.

```python
import numpy as np
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')
%matplotlib inline
```

Now let's plot $f(x)=x^6-x-1$.

```python
x = np.arange(-1,1.5,0.05)
f = lambda x: x**6-x-1
plt.plot(x, f(x), color = 'black', ms = 3);
```

There seems to be a root somewhere on the interval $[-1.0, -0.5]$. Let's see if we can implement our algorithms to approximate its value. 

# 1. Review of the Bisection Method

Now that we know how to use lambda functions we can update the way we wrote our `bisection_method` function. We can now pass the function in as a parameter.

```python
def bisection_method(f, a, b, iterations = 25, tol = 1e-6):
    """
    Parameters
    ----------
    f               : Function for which we are searching for a solution f(x)=0
    a, b            : The endpoints of the interval
    iterations = 25 : Set the maximum number of iterations for the loop
    tol = 1e-6      : Set the level of tolerance for the distance between x_i and the midpoint
    
    Returns
    -------
    The approximation for the root (if found)
    """
    
    if f(a)*f(b) >= 0:
        print("A root is not guaranteed in this interval.")
        return 0
    
    for n in range(iterations):
        c = (a + b)/2.0
                          
        if f(a)*f(c) > 0:
            a = c
        else:
            b = c
            
        if abs(b - a) < tol:
            break
        elif n == iterations - 1:
            print("\nLevel of tolerance not reached after 25 iterations.")
            break
    print('Found solution after', n, 'iterations.')
    return c
```

<!-- BEGIN QUESTION -->

**Question 3.** Use the `bisection_method` function to approximate the solution to $f(x)=x^6-x-1$. 

```python
p = ...
Dp = ...
root = bisection_method(..., -1, -0.5)
print(root)
```

<!-- END QUESTION -->

## 2. Newton's Method

[Newton's method](https://en.wikipedia.org/wiki/Newton%27s_method) is a root finding method that uses linear approximation. In particular, we make an initial guess ($x_0$) for the root of an equation ($f(x)=0$), compute the linear approximation of $f(x)$ at $x_0$ and then find the $x-$intercept of the linear approximation. The algorithm we learned in class is 

$$x_{n+1}=x_n-\frac{f\left(x_n\right)}{f'\left(x_n\right)}$$

which converges (potentially) to an approximate solution to the equation $f(x)=0$.

<!-- BEGIN QUESTION -->

**Question 4.** In class and in the previous lab we discussed situations in which Newton's method would fail to converge, Compare Newton's method to the bisection method. What are the differences? What are the advantages/disadvantages of Newton's method?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Below is a function that uses Newton's method to approximate a solution to $f(x)=0$. Read through the description and code. If you have questions you can ask another student in your pod or ask one of your instructors. 

```python
def newtons_method(f, Df, x0, iterations = 25, tol = 1e-6):
    """
    Parameters
    ----------
    f               : Function for which we are searching for a solution f(x)=0.
    Df              : Derivative of f(x)
    x0              : Initial guess for a solution f(x)=0
    iterations = 25 : Set the aximum number of iterations for the loop
    tol = 1e-6      : Set the level of tolerance for a stopping criteria abs(f(x)) < tol

    Returns
    -------
    The approximation of the root (if found)
    """
    xn = x0
    for n in range(0, iterations):
        fxn = f(xn)
        if abs(fxn) < tol:
            print('Found solution after', n, 'iterations.')
            return xn
        
        Dfxn = Df(xn)
        
        if Dfxn == 0:
            print('Zero derivative. No solution found.')
            return None
        
        xn = xn - fxn/Dfxn
    print('Exceeded maximum iterations. No solution found.')
    return None
```

Now that you've read through the description of the function and the looked over the code, let's use `newtons_method`. Below you will find a use case for the function:

**Example 1.** Use the `newtons_method` function to approximate the solution to the function $x^2-x-1$.

```
p = lambda x : x**2-x-1
Dp = lambda x: 2*x-1
newtons_method(f, Df, 1)

Found solution after 5 iterations.
1.618033988749989
```

<!-- BEGIN QUESTION -->

**Question 5.** Use the `newtons_method` function to approximate the solution to the function $x^2-x-1$.

```python
p = ...
Dp = ...
root = newtons_method(..., ..., 1)
#statement = "The root is {}."
#statement.format(root)
print(root)
```

<!-- END QUESTION -->

### Supergolden Ratio

In mathematics, two quantities are in the supergolden ratio if the quotient of the larger number divided by the smaller one is equal to

$${\displaystyle \psi ={\frac {1+{\sqrt[{3}]{\frac {29+3{\sqrt {93}}}{2}}}+{\sqrt[{3}]{\frac {29-3{\sqrt {93}}}{2}}}}{3}}}$$

which is the only real solution to the equation $x^{3}=x^{2}+1$. [(Wikipedia)](https://en.wikipedia.org/wiki/Supergolden_ratio)

Let's test our function `newtons_method` on the polynomial $f(x)=x^3-x^2-1$ to approximate the super golden ratio.

<!-- BEGIN QUESTION -->

**Question 6.** Use the `newtons_method` function to approximate the solution to the function $x^3-x^2-1$.

```python
p = ...
Dp = ...
root = newtons_method(..., ..., 1)
print(root)
```

<!-- END QUESTION -->

Five iterations is really fast! During class Mr. Rash mentioned how the limit as $n \to \infty$ (where $n$ is the number of iterations of Newton's method) of the ratio of the most recent error to the square of the previous error approaches a constant. More specifically

$$\lim\limits_{n \to \infty} \frac{\left|\alpha-x_{n+1}\right|}{\left|\alpha-x_n \right|^2}=\frac{-f''(\alpha)}{2f'(\alpha)}$$

Since we know the real solution to $x^3-x^2-1=0$ we can make a table of values for the 5 iterations it took our algorithm to approximate the solution. 

First lets get the value of the only real solution (i.e. the supergolden ration) stored in a python variable. You may find it useful to use numpy for your calculations.

**Question 7.** Write a python expression to calculate the value of the supergolden ratio.

```python
sg_ratio = ...
sg_ratio
```

```python
grader.check("q7")
```

Now we can use our value of alpha ($\alpha$) to make our table.

**Note:** The value for `sg_ratio` is our value of the "true" root alpha ($\alpha$).

Run the cell below to generate a list of $x-$values from Newton's method.

```python
def newtons_method_values(f, Df, x0, iterations = 25, tol = 1e-6):
    """
    Parameters
    ----------
    f               : Function for which we are searching for a solution f(x)=0.
    Df              : Derivative of f(x)
    x0              : Initial guess for a solution f(x)=0
    iterations = 25 : Set the aximum number of iterations for the loop
    tol = 1e-6      : Set the level of tolerance for a stopping criteria abs(f(x)) < tol

    Returns
    -------
    The approximation of the root (if found)
    """
    xis = []
    xn = x0
    for n in range(0, iterations):
        fxn = f(xn)
        if abs(fxn) < tol:
            return xis
        
        Dfxn = Df(xn)
        
        if Dfxn == 0:
            print('Zero derivative. No solution found.')
            return None
        
        xn = xn - fxn/Dfxn
        xis.append(xn)
    print('Exceeded maximum iterations. No solution found.')
    return xis
```

Now we can make a table similar to the one in your textbook *Elementary Numerical Analysis* on Page 81.

Run the cell below.

```python
p = lambda x : x**3-x**2-1
Dp = lambda x: 3*x**2-2*x
xis = newtons_method_values(p, Dp, 1)
print("{:<8} {:<25} {:<25} {:<30}".format("i", "x_i", "alpha-x_i", "|alpha-x_{n+1}|/|alpha-x_{n}|^2"))
for i in range(len(xis)):
    if i == 0:
        print("{:<8} {:<25} {:<25}".format(i, xis[i], abs(xis[i]-sg_ratio)))
    elif i ==  3:
        print("{:<8} {:<25} {:<25} {:<30}".format(i, xis[i], abs(xis[i]-sg_ratio), "a."))
    elif i == 4:
        print("{:<8} {:<25} {:<25}".format(i, "b.", "c."))
    else:
        print("{:<8} {:<25} {:<25} {:<30}".format(i, xis[i],abs(xis[i]-sg_ratio), abs(xis[i+1]-sg_ratio)/(xis[i]-sg_ratio)**2))
```

<!-- BEGIN QUESTION -->

**Question 8.** Calculate by hand the values for **a**, **b**, and **c** (at least 6 decimal places). Then Type them in the cell below.

**Note:** You can use python, a spreadsheet or your calculator.

```python
a = ...
b = ...
c = ...
```

```python
grader.check("q8")
```

<!-- END QUESTION -->

**Question 9.** Calculate the value of $\displaystyle \frac{-f''(\alpha)}{2f'(\alpha)}$. You are allowed to use more than one line of code to compute the value.

**Hint:** You may want to define another lambda function. 

```python
lim_as_goes_to_infinity = ...
lim_as_n_goes_to_infinity
```

```python
grader.check("q9")
```

<!-- BEGIN QUESTION -->

**Question 10.** How close is the value from **Question 9.** to your answer for part **a.**

```python
error = ...
error
```

<!-- END QUESTION -->

Let's change the level of tolerance for our `newtons_method` function to `1e-10`.

Run the cell below.

```python
p = lambda x : x**3-x**2-1
Dp = lambda x: 3*x**2-2*x
root = newtons_method(p, Dp, 1, tol = 1e-10)
print(root)
print("The error between Newtons method and the 'exact value' of the solution to p(x) is", abs(sg_ratio-root))
```

<!-- BEGIN QUESTION -->

**Question 11.** How many iterations of the Bisection method starting with the interval of $[1,2]$ can achieve the same accuracy? Explain how you determined your value.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 12.** We discussed cases in which Newton's method would fail to converge (i.e. diverge). For example, if the tangent line at the root is vertical as is the case for $f(x)=x^{1/3}$. Explain why this would be an issue for Newton's method and not for the Bisection method.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Secant Method

The Secant method is very similar to the bisection method. The bisection method divides each interval by choosing a midpoint while the secant method divides each interval by a secant line. The secant line always converges to a root (provided that $f(x)$ is continuous on $[a,b]$ and $f(a) \cdot f(b)<0$). The algorithm for the secant method is 

$$x_{n+1}=x_n-f(x_n) \cdot \frac{x_n-x_{n-1}}{f(x_n)-f(x_{n-1})}$$

Let's see if we can write a python function to implement the Secant method.

**Question 12.** Fill in the ellipses for the `secant_method` function. Then test your function on the polynomial $f(x)=x^3-x^2-1$.

```python
def secant_method(f, a, b, iterations = 25):
    """

    Parameters
    ----------
    f               : Function for which we are searching for a solution f(x)=0.
    a, b            : Endpoints of the interval
    iterations = 25 : Set the aximum number of iterations for the loop

    Returns
    -------
    The approximation of the root (if found)
    """
 if f(a)*f(b) >= 0:
        print("Secant method fails.")
        return None
    a_n = a
    b_n = b
    for n in range(1, iterations + 1):
        m_n = a_n - f(...)*(... - ...)/(f(...) - f(...))
        f_m_n = f(m_n)
        
        if f(a_n)*f_m_n < 0:
            a_n = ...
            b_n = ...
        elif f(b_n)*f_m_n < 0:
            a_n = ...
            b_n = ...
        elif f_m_n == 0:
            print("Found a solution")
            return m_n
        else:
            print("Secant method fails.")
            return None
    return ... - f(...)*(... - ...)/(f(...) - f(...))
    
 = 
f = lambda x: x**3-x**2-1
root = secant_method(f, 1, 2)
print(root)
```

```python
grader.check("q12")
```

<!-- BEGIN QUESTION -->

**Question 13.** How close is the value from **Question 12.** to your the value of the supergolden ratio (`sg_ratio`)?

```python
f = lambda x: x**3-x**2-1
root = ...
error = ...
error
```

<!-- END QUESTION -->

# 4. Comparing Methods

<!-- BEGIN QUESTION -->

**Question 13.** Consider the function $$f_1(x)=x^4-7.5x^3+10.56x^2+31.744x-68.8128$$, which has a root at $x=3.2$.

**a.** Use all three root finding methods to find the root at $x=3.2$. Use the following specifications:

- Tolerance for each method: $10^{-7}$
- Maximum number of iterations for each method: 25
- Initial interval/points for the Bisection and Secant methods: $[3.0, 3.5]$
- Initial point for Newton's method: 3.0

**b.** Now consider $$f_2(x)=x^2-1.1x-6.72$$, which also has a root at $x=3.2$. Using the same specifications as above, find the root at $x=3.2$ using all three root finding methods.

**Note:** You are allowed to use python, a spreadsheet or your calculator.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 14.** Compare and contrast our various root-finding methods. What did you learn about their  strengths and weaknesses? Which seem to work “best”? Support your statements with evidence. Be clear and concise in your comparisons.

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: None)
```
>>> f(1)
-1
```
```
>>> f(0)
-1
```
```
>>> f(10)
999989
```


Tests for q12 (points: None)
```
>>> np.isclose(1.4655712318713536, root)
True
```


Tests for q2 (points: None)
```
>>> Df(1)
5
```
```
>>> Df(0)
-1
```
```
>>> Df(10)
599999
```


Tests for q7 (points: None)
```
>>> np.isclose(1.4655712318767682, sg_ratio)
True
```


Tests for q8 (points: None)
```
>>> np.isclose(0.9665207024404882, a)
True
```
```
>>> np.isclose(1.4655713749070918, b)
True
```
```
>>> np.isclose(1.4303032358498058e-07, c)
True
```


Tests for q9 (points: None)
```
>>> np.isclose(0.9670208838033376, lim_as_n_goes_to_infinity)
True
```
