```python
# Initialize Otter
import otter
grader = otter.Notebook("lab12.ipynb")
```

# Lab 12: Differential Equations

Welcome to Lab 12! Throughout the course you will complete a lab assignments like this one. You can't learn technical subjects without hands-on practice, so labs are an important part of the course. 

Collaborating on labs is more than okay -- it's encouraged. You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too $-$ the best way to solidify your knowledge of a subject is to explain it. You should not just copy someone else's answers, but rather work together to gain understanding of the task you need to complete.

In today's lab, you will be asked to use your python programming skills to write code to approximate solutions to differential equations.

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** Tuesday, May 17, 2022 at 11:59 pm

**Collaboration Policy:** Labs are a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

Run the cell below to import the required modules.

```python
import numpy as np
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')
%matplotlib inline
```

Differential equations are relationships between a function and its derivatives, and they are used to model systems in every engineering and science field. For example, a simple differential equation relates the acceleration of a car with its position. Unlike differentiation where analytic solutions can usually be computed, in general finding exact solutions to differential equations is very hard. Therefore, numerical solutions are critical to making these equations useful for designing and understanding engineering and science systems. [(Numerical Methods with Python)](https://pythonnumericalmethods.berkeley.edu/notebooks/Index.html)

# 1. Euler's Method

Let's write a function to implement Euler's method. Instead of doing it all at one time we will complete it in small sections. Then we will put it all together. For the purpose of this activity, let's suppose that our initial value problem is 

$$\frac{dy}{dx}=-y+\sin(x)\text{, }y(0)=1$$

We will use Euler's method to approximate the value of the function in the interval $[0,10]$ with 100 points.

**Question 1.** Assign the values for the endpoints of the interval, the initial value, and the change in $x$.

```python
x0 = ...
y0 = ...
xn = ...
n = ...
deltax = ...
```

```python
grader.check("q1")
```

We next want to create a vector of length 101 whose entries are our $x-$values. To do this, we enter the following code

```python
x = np.linspace(x0, xn, n+1)
x
```

Now the computer has a vector called `x`, given by 

```
x = [0.0, 0.1, 0.2, ..., 10.]
```

The next thing we want to do generate the $y-$values. The first thing we do is create an array to store the $y-$values.

```python
y = np.zeros ([n+1])
y
```

This creates an array `y` with n entries that are all 0.0.

**Question 2.** The next thing we want to do is fix this so that the entries are the estimates given by Euler's method. To do this, we use a `for` loop.

```python
y[0] = y0 
for i in range(1, n+1): 
    y[i] = ...
y
```

```python
grader.check("q2")
```

We can also use a for loop to print out the data that we’ve generated. Let's only print the first 10 ordered pairs.

```python
for i in range(101): 
    print (i, "\t", round(x[i], 2), "\t", y[i])
```

Finally, to make a graph of our data, we can use the plot function from `matplotlib`.

```python
plt.scatter(x, y, color = 'black', s = 3) 
plt.xlabel("x") 
plt.ylabel("y")
plt.title("Approximate Solution with Eulers Method");
```

# 2. Runge-Kutta

The Runge-Kutta method is a mathematical algorithm used to solve systems of ordinary differential equations (ODEs) by using a trial step at the midpoint of an interval to cancel out lower-order error terms. The second-order formula (sometimes known as RK2) is

$k_1=h\cdot f(x_n,y_n)$

$k_2=h\cdot f(x_n+\frac{1}{2}h, y_n+\frac{1}{2}k_1)$

$y_{n+1}=y_n+k_2$

and the fourth order formula (sometimes known as RK4) is 

$k_1=h\cdot f(x_n,y_n)$

$k_2=h\cdot f(x_n+\frac{1}{2}h, y_n+\frac{1}{2}k_1)$

$k_3=h\cdot f(x_n+h, y_n+k_3)$

$k_4=h\cdot f(x_n+h, y_n+k_3)$

$y_{n+1}=y_n+\frac{1}{6}k_1+\frac{1}{3}k_2+\frac{1}{3}k_3+\frac{1}{6}k_4 = \frac{h}{6}(k_1+2k_2+2k_3+k_4)$

<!-- BEGIN QUESTION -->

**Question 2.** Write a function to implement the fourth order Runge-Kutta method.

```python
    ...
```

<!-- END QUESTION -->

Use your python function from **Question 2** on 

$$\frac{dy}{dx}=\frac{1}{1+x^2}-2y^2$$

with initial condition $(0,1)$, $h=0.05$ and 20 steps to estimate $y(1.0)$.

Print the $x$, $y$ coordinates at each step. Your output should look like this

```
0 	  0      1
1 	  0.05   0.954574885889361
2 	  0.1    0.9168244152086702
3 	  0.15   0.8849543164090213
4 	  0.2    0.8575987608041347
5 	  0.25   0.8337068961117725
6 	  0.3    0.8124658691548322
7 	  0.35   0.7932465576266018
8 	  0.4    0.7755638210856183
9 	  0.45   0.7590463362380451
10 	 0.5    0.7434130154950842
11 	 0.55   0.7284541649529332
12 	 0.6    0.7140162267094241
13 	 0.65   0.6999893545301911
14 	 0.7    0.686297304791557
15 	 0.75   0.6728892575011016
16 	 0.8    0.659733259237782
17 	 0.85   0.6468110276680971
18 	 0.9    0.6341138910327994
19 	 0.95   0.6216396635803484
20 	 1.0    0.6093902829524289
```

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: None)
```
>>> x0 == 0
True
```
```
>>> y0 == 1
True
```
```
>>> xn == 10
True
```
```
>>> deltax == 0.1
True
```


Tests for q2 (points: None)
```
>>> len(y)
101
```
```
>>> sum(y)
27.12670698196994
```
