```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

# Lab 06: More Root Finding and Fixed Point Iteration

Welcome to Lab 06! Throughout the course you will complete a lab assignments like this one. You can't learn technical subjects without hands-on practice, so labs are an important part of the course.

Collaborating on labs is more than okay -- it's encouraged. You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** just copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

In today's lab, you'll learn more about root finding algorithms and fixed-point iteration. 

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** Tuesday, March 15, 2022 at 11:59 pm

**Collaboration Policy:** Labs are a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

# 0. Importing Modules

In this lab we will need the following modules:

- `NumPy`

- `math`

- `matplotlib`


Run the cell below to import the required modules.

```python
import numpy as np
import math
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')
%matplotlib inline
```

**Question 1.** Using `np.sin` we can find $y-$values for $f(x)=\sin (x)$. Use `np.sin` to find the sine of $\frac{\pi}{3}$.

```python
sin_pi_over_2 = ...
sin_pi_over_2
```

```python
grader.check("q1")
```

# 1. Review of the Root Finding Methods

Below are example Python functions for each of the root finding methods we've studied in our course: 

* Bisection method

* Secant method

* Newton's method

```python
def bisection_method(f, a, b, iterations = 25, tol = 1e-6):
    """
    Parameters
    ----------
    f               : Function for which we are searching for a solution f(x)=0.
    a, b            : The endpoints of the interval.
    iterations = 25 : Set the maximum number of iterations for the loop.
    tol = 1e-6      : Set the level of tolerance for a stopping criteria.
    
    Returns
    -------
    The approximation for the root (if found).
    """
    
    """
    Intermediate Value Theorem check.
    """
    if f(a)*f(b) >= 0:
        print("A root is not guaranteed in this interval.")
        return 0
    
    """
    For loop for the algorithm.
    """
    for n in range(iterations):
        """
        Bisection method
        """
        c = (a + b)/2.0
                          
        if f(a)*f(c) > 0:
            a = c
        else:
            b = c
            
        """
        Error check and iterations check.
        """
        if abs(b - a) < tol:
            break
        elif n == iterations - 1:
            print("\nLevel of tolerance not reached after 25 iterations.")
            break
    
    print('Found solution after', n, 'iterations.')
    return c

def secant_method(f, a, b, iterations = 25, tol = 1e-6):
    """

    Parameters
    ----------
    f               : Function for which we are searching for a solution f(x)=0.
    a, b            : Endpoints of the interval.
    iterations = 25 : Set the aximum number of iterations for the loop.
    tol = 1e-6      : Set the level of tolerance for a stopping criteria.

    Returns
    -------
    The approximation of the root (if found).
    """
    
    """
    Initialize the err, f(a), and iteration counter n.
    """
    err = 1
    fa = f(a)
    n = 0
    
    """
    While loop for algorithm
    """
    while abs(err) > tol and n <= iterations:
        
        """
        Update iteration counter.
        """
        n += 1
        fb = f(b)
        
        """
        Check to see if there will be division by 0.
        """
        if fa - fb == 0:
            print('f(a)-f(b)=0: Division by zero is undefined')
            return None
        
        """
        Secant method and error check.
        """
        c = b - fb*(b-a)/(fb-fa)
        err = c - b
        
        a = b 
        b = c
        fa = fb
        
        """
        Check the number of iterations.
        """
        if n >= iterations:
            print('Exceeded maximum iterations', n, 'No solution found.')
            return None
    
    print('Found solution after', n, 'iterations.')
    return c

def newtons_method(f, Df, x0, iterations = 25, tol = 1e-6):
    """
    Parameters
    ----------
    f               : Function for which we are searching for a solution f(x)=0.
    Df              : Derivative of f(x).
    x0              : Initial guess for a solution f(x)=0.
    iterations = 25 : Set the maximum number of iterations for the loop.
    tol = 1e-6      : Set the level of tolerance between successive x-values for a stopping criteria.

    Returns
    -------
    The approximation of the root (if found).
    """
    
    """
    Initialize the err, x, and iteration counter n.
    """
    err = 1
    x = x0
    n = 0
    
    """
    While loop for algorithm.
    """
    while abs(err) > tol and n <= iterations:
        fxn = f(x)
        Dfxn = Df(x)
        
        """
        Check to see if the derivative at x0 equals 0.
        """
        if Df(x0) == 0:
            print('Zero derivative. No solution found.')
            return None
        
        """
        Newtons method and error check
        """
        xn = x - fxn/Dfxn
        err = xn - x
        x = xn
        
        """
        Update iteration counter.
        """
        n += 1
        
        """
        Check the number of iterations.
        """
        if n >= iterations:
            print('Exceeded maximum iterations', n, 'No solution found.')
            return None
    
    print('Found solution after', n, 'iterations.')
    return xn
```

<!-- BEGIN QUESTION -->

**Question 2.** The function $f(x)=1-x+\sin⁡(x$) has a root in the interval $[1,3]$. Plot the function on the interval.

**Hint:** It might help to use a lambda function and `numpy`'s version of the sine function (`np.sin`).

```python
...

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.** How many iterations of bisection would be needed in order to approximate the root of the function in **Question 2** with error no more than $5 \cdot 10^{-8}$? Justify your answer.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

A trough of length $L$ has a cross section in the shape of a semicircle with radius $r$ (see figure below). 

<center><img src="images/watertrough.jpg" width="600" height="700"/></center>

When filled with water to within a distance $h$ of the top, the volume $V$ of water is

$$V=L \left[ \frac{1}{2}\pi r^2-r^2\arcsin \left(\frac{h}{r} \right)-h\sqrt{r^2-h^2}\right]$$

**Note:** The arcsine function can be accessed fron the `numpy` module using `np.arcsin`.

Suppose $L=15$ ft, $r=2.0$ ft, and $V=75.0$ ft$^3$. Find the depth of water in the trough to within $0.0001$ ft. Use each of the three functions (`bisection_method`, `secant_method`, and `newtons_method`. Select your own reasonable initial condition(s). 

**Hint:** It may help to notice that $h$ is necessarily bounded between $0$ and $r$.

<!-- BEGIN QUESTION -->

**Question 4.**  Use the `bisection_method` function.

**Hint:** It will be helpful to use a lambda function for $f(x)$.

```python
f = lambda ...
root = ...
root
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5.**  Use the `secant_method` function.

**Hint:** It will be helpful to use a lambda function for $f(x)$.

```python
f = lambda ...
root = 
root
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6.**  Use the `newtons_method` function.

**Hint:** It will be helpful to use a lambda functions for $f(x)$ and $f'(x)$. 

**Warning:** Check your derivative with a classmate before implementing Newton's method.

```python
f = lambda ...
df = lambda ...
root = ...
root
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7.**  Compare and contrast what happens in **Questions 4**, **5**, and **6**. Especially address the “speed” of convergence toward the root. Are you surprised by your results? Explain.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8.** In class we looked at some examples of fixed-point iterations (see the Fixed-Point Iteration tab in One Note). Consider the sequence $x_{n+1}=g\left(x_n \right)$ for each $g(x)$ below.

**a.** $\displaystyle g_1(x)=\frac{5}{x}$

**b.** $\displaystyle g_2(x) = 1+x-\frac{x^2}{5}$

**c.** $\displaystyle g_3(x) = \frac{1}{2}\left(x + \frac{5}{x} \right)$

**d.** $\displaystyle g_4(x) = 5+x-x^2$

Verify that the $\sqrt 5$ is a fixed-point solution for each function.

**Note:** In class we looked at some examples (see the Fixed-Point Iteration tab in One Note)

```python
g_1 = lambda ...
g_2 = lambda ...
g_3 = lambda ...
g_4 = lambda ...
print("The square root of 5 is", np.sqrt(5))
print("g_1 @ square root 5 =", g_1(np.sqrt(5)))
print("g_2 @ square root 5 =", g_2(np.sqrt(5)))
print("g_3 @ square root 5 =", g_3(np.sqrt(5)))
print("g_1 @ square root 5 =", g_4(np.sqrt(5)))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 9.** Use the results from Tuesday's class to determine which functions $\left(g_1, g_2, g_3, \right .$ or $\left . g_4\right)$ converge. Justify your answers.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 10.** Each of the functions from **Question 9** are repeated below. If a function converges to $\sqrt 5$, how "fast" (qualitatively speaking), does it converge? Feel free to use a code cell in this notebook, another programming language, your calculator, or a spreadsheet to do your calculations. If a function does not converge indicate as such in the text cell.

<!-- BEGIN QUESTION -->

**Question 10 a.**  $\displaystyle g(x)=\frac{5}{x}$

**Note:** You can use the cell below to run code.

```python

```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10 b.**  $\displaystyle g(x) = 1+x-\frac{x^2}{5}$

**Note:** You can use the cell below to run code.

```python

```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10 c.**  $\displaystyle g(x) = \frac{1}{2}\left(x + \frac{5}{x} \right)$

**Note:** You can use the cell below to run code.

```python

```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10 d.**  $\displaystyle g(x) = 5+x-x^2$

**Note:** You can use the cell below to run code.

```python

```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 11.** Observe the table below for $g(x)=1+x-\frac{1}{5}x^2$, (from your textbook on page 104) which converges to the $\sqrt 5$. Also, observe that $r_n$ is approaching $g'(\alpha)$.

- $n$

- $x_n$

- $\alpha -x_n$

- $r_n=\frac{\alpha - x_n}{\alpha - x_{n-1}}$


|n    | x_n          | alpha-x_n   | r_n       |
|-----|--------------|-------------|-----------|
| 0   | 2.5          | -2.64E-1    |           |
| 1   | 2.25         | -1.39E-2    | 0.0528    |
| 2   | 2.2375       | -1.43E-3    | 0.1028    |
| 3   | 2.23621875   | -1.51E-4    | 0.1053    |
| 4   | 2.23608389   | -1.59E-5    | 0.1055    |
| 5   | 2.23606966   | -1.68E-6    | 0.1056    |
| 6   | 2.23606815   | -1.77E-7    | 0.1056    |
| 7   | 2.23606800   | -1.87E-8    | 0.1056    |

For any other function(s) from **Question 9** that converge construct a similar table.

**Notes:** You can use the cell below to run code. 

```python

```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 12.** Verify from table(s) that $r_n$ approaches $g'(\alpha)$.

```python

```

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: None)
```
>>> sin_pi_over_2
0.8660254037844386
```
